---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Insert exercise title here

```yaml
type: VideoExercise
key: 975340a4d6
xp: 50
```

`@projector_key`
a381b9470e2bca79e6602a2692d26245

---

## Insert exercise title here

```yaml
type: VideoExercise
key: 9a8deeb367
xp: 50
```

`@projector_key`
331014cfbd7a53d7c05126c996d02831
